# Начало работы 
Темы: 
1. Как печатать 
2. Как назначать переменные 
3. Как комментировать 
4. Синтаксис для базовой математики

## Как печатать

В Julia для печати используется функция `print()`. 
Для печати с переходом на следующую строку попробуйте `println()`

```julia

```

```julia
println("- Я буду изучать Julia!") 
print("- Молодец, ")
print("иди лучше посуду помой")
```

```julia

```

## Как считывать с клавиатуры

Для перехвата строки вводимой на клавиатуре используется `readline`. При выполнении появляется специальное поле, куда можно ввести строку

```julia
println("Как Вас зовут?") 
name = readline()
println("Привет, $name!") 
```

```julia
println("Ведите два числа") 
x = readline()
y = readline()
z = x*y
```

```julia
typeof(z)
```

```julia
println("Ведите два числа") 
x = Meta.parse( readline() )
y = Meta.parse( readline() )
z = x*y
```

```julia
typeof(z)
```

## Как назначить переменные 

Все, что нам нужно, это имя переменной, значение и знак равенства! Юлия сама определит для нас типы.

```julia
my_answer = 42
typeof(my_answer)
```

```julia
my_pi = 3.14159
typeof(my_pi)
```

```julia
a, b = 5, 7
```

```julia
a+b
```

```julia
x = readline()
```

```julia
ᛤ = "cealc" # поддержка Юникод символов
typeof(ᛤ)
```

```julia
औ = 1 # вы ведь всегда хотели попробовать создать переменную на Деванагари
```

```julia
typeof(औ)
```

Чтобы ввести греческий символ используйте LaTeX команду с последующим нажатием `Tab` 

```julia
δ = 'δ' # \delta + <tab> --> δ
typeof(δ)
```

После присвоения значения переменной мы можем без проблем переназначить значение другого типа этой переменной.

```julia
my_pi = "примерно 3"
```

```julia
typeof(my_pi)
```

```julia
ᛃ = 0
ᛥ = -1
```

```julia
ᛃ + ᛥ == -औ
```

## Комментарии

```julia
# Вы можете оставить комментарии в одной строке, используя клавишу фунта / хеша
```

```julia
#=

Для многострочных комментариев 
используйте последовательность '#= =#'.

=#
```

## Синтаксис для базовой матматики

```julia
sum = 3 + 7
```

```julia
difference = 10 - 3
```

```julia
product = 20 * 5
```

```julia
quotient = 100 / 10 
```

```julia
power = 10 ^ 2
```

```julia
modulus = 101 % 2
```

```julia
sqrt(2) # square root
```

```julia
√2 # можно использовать Unicode с помощью: \sqrt + TAB
```

Обратите внимание, что деление двух целых чисел дает число с плавающей запятой. Здесь могут быть полезны два дополнительных оператора:

```julia
10 / 6
```

```julia
10 ÷ 6 # введите \div и нажмите TAB или используйте функцию `div` для целочисленного деления
```

```julia
div(10, 6)
```

```julia
10 // 6 # представление в виде дроби
```

### Числа: несколько способов написания числа "сорок два"

```julia
fortytwos = (42, 42.0, 4.20e1, 4.20f1, 84//2, 0x2a)
```

```julia
for x in fortytwos
    show(x)
    println("\tisa $(typeof(x))")
end
```

### Побитовая арифметика

```julia
0x2a & 0x70 # AND
```

```julia
0x2a | 0x70 # OR
```

```julia
42 & 112
```

```julia
0b0010 << 2 # == 0b1000
```

```julia
?<<
```

### Логические операторы

```julia
false && true # AND
```

```julia
false || true # OR
```

Заметьте, что это так называемые "операторы короткого замыкания"

```julia
x = -42
x > 0 || error("x must be positive")
```

### Сравнение

```julia
1 == 1.0 # Равны ли
```

```julia
1 === 1.0 # программно идентичны
```

```julia
0.2 == .2
```

```julia
3 < π
```

```julia
1 <= 1
```

```julia
.1 + .2
```

```julia
.1 + .2 == .3
```

```julia
.1 + .2 ≈ .3 # \approx + TAB
```

Цепочка сравнений

Попробуйте вставить скобки вокруг одного из этих сравнений

```julia
2 == 2.0 == 0x02
```

```julia
x = 42
0 < x < 100 || error("x must be between 0 and 100")
```

## Операторы обновления 

Все вышеприведенные арифметические операторы инфикса могут использоваться как операторы «обновления» в сочетании с присваиванием:

```julia
y += 1
```

```julia
# Это тоже самое что и:
y = y + 1
```

```julia
typeof(y)
```

Обратите внимание, что это просто повторное использование такого же имени для нового значения. Это означает, что тип может даже измениться:

```julia
y /= 2
```

```julia
typeof(y)
```

### Упражнения

#### 1.1 
Найдите в документации информацию по функции конверт `convert`.

```julia

```

#### 1.2 
Создайте переменную `days` и присвойте ей значение `365`. Конвертируйте `days` в тип float и присвойте полученное значение переменной `days_float`

```julia

```

<!-- grade_id: cell-715f78016beb0489, points: 1, grade: true, solution: false -->
```julia
@assert days == 365
@assert days_float == 365.0

```

#### 1.3 
Посмотрите, что происходит, когда вы выполняете

```julia
convert(Int64, "1")
```
и

```julia
parse(Int64, "1")
```

```julia

```