# Основы: Арифметика, Присвоение и Строки 

Темы: 
1. Инфиксные арифметические операторы 
2. Числовые литералы 
3. Сравнение 
4. Присвоение 
5. Строки

## Арифметика должно быть вам в значительной степени знакома

```julia
3 + 7 # addition
```

```julia
10 - 3 # subtraction
```

```julia
20 * 5 # multiplication
```

```julia
100 / 10 # division
```

```julia
10 ^ 2 # exponentiation
```

```julia
101 % 2 # remainder (остаток)
```

```julia
sqrt(2) # square root
```

```julia
√2 # можно использовать Unicode с помощью: \sqrt + TAB
```

Обратите внимание, что деление двух целых чисел дает число с плавающей запятой. Здесь могут быть полезны два дополнительных оператора:

```julia
10 / 6
```

```julia
10 ÷ 6 # введите \div и нажмите TAB или используйте функцию `div` для целочисленного деления
```

```julia
div(10, 6)
```

```julia
10 // 6 # представление в виде дроби
```

### Числа: несколько способов написания числа "сорок два"

```julia
fortytwos = (42, 42.0, 4.20e1, 4.20f1, 84//2, 0x2a)
```

```julia
for x in fortytwos
    show(x)
    println("\tisa $(typeof(x))")
end
```

### Побитовая арифметика

```julia
0x2a & 0x70 # AND
```

```julia
0x2a | 0x70 # OR
```

```julia
42 & 112
```

```julia
0b0010 << 2 # == 0b1000
```

```julia
?<<
```

### Логические операторы

```julia
false && true # AND
```

```julia
false || true # OR
```

Заметьте, что это так называемые "операторы короткого замыкания"

```julia
x = -42
x > 0 || error("x must be positive")
```

### Сравнение

```julia
1 == 1.0 # Равны ли
```

```julia
1 === 1.0 # программно идентичны
```

```julia
0.2 == .2
```

```julia
3 < π
```

```julia
1 <= 1
```

```julia
.1 + .2
```

```julia
.1 + .2 == .3
```

```julia
.1 + .2 ≈ .3 # \approx + TAB
```

Цепочка сравнений

Попробуйте вставить скобки вокруг одного из этих сравнений

```julia
2 == 2.0 == 0x02
```

```julia
x = 42
0 < x < 100 || error("x must be between 0 and 100")
```

# Присваивание

Инициализация в Julia выполняется с помощью одиночного `=`. Этот оператор просто связывает имя (слева) со значением (справа).

```julia
x = 1 # использовать имя `x` для значения `1` (создали переменную икс инициализировав её единицей)
```

```julia
y = x # использовать имя `y` для значения на которое ссылается `x` 
```

```julia
x = "hello!" # Решите, как вам лучше использовать имя `x`
```

```julia
y # всё еще равен 1
```

```julia
ϵ = eps(1.0) # можете создавать переменные используя unicode
```

```julia
5ϵ # в таких случаях знак * можно опустить
```

## Операторы обновления 

Все вышеприведенные арифметические операторы инфикса могут использоваться как операторы «обновления» в сочетании с присваиванием:

```julia
y += 1
```

```julia
# Это тоже самое что и:
y = y + 1
```

```julia
typeof(y)
```

Обратите внимание, что это просто повторное использование такого же имени для нового значения. Это означает, что тип может даже измениться:

```julia
y /= 2
```

```julia
typeof(y)
```

# Строки

```julia
s1 = "Это строка."
```

```julia
s2 = """Это тоже строка. """
```

```julia
"Здесь будет "error" потому что не понятно, где начинается какая строка "
```

```julia
"""Гляньте-ка, нет больше "error"!!! """
```

```julia
println("""Еще одна приятная вещь в строковых литералах
    с тройными кавычками заключается в том, что они игнорируют
    начальные отступы, что хорошо для длинных строк в реальном коде.
    Попробуйте здесь чего-нибудь поменять!""")
```

Строки не пишутся с одиночными `'` кавычками - их используют для одиночного символа:

```julia
first(s1)
```

```julia
'⊂'
```

```julia
'If you try writing a string in single-quotes, you will get an error'
```

## Строковая интерполяция

Вы можете использовать знак доллара внутри строки, чтобы отобразить выражение Julia внутри строки - либо одну переменную, либо более сложное выражение:

```julia
name = "Вася"
num_fingers = 10
num_toes = 10
println("Привет, я $name.")
println("У меня на руках $num_fingers пальцев и на ногах $num_toes.")
println("Итого $(num_fingers+num_toes) . Согласен, пример так себе...")
```

```julia

```

<!-- grade_id: cell-715f78016beb0489, points: 1, grade: true, solution: false -->
```julia
@assert days == 365
@assert days_float == 365.0

```

Output:
```
Stacktrace:
 [1] top-level scope at In[23]:1
```