# Строки

Темы: 
1. Как получить строку 
2. Строковая интерполяция 
3. Конкатенация строк

## Как получить строку

Просто! Нужно заключить набор символов в кавычки " " или """ """!

```julia
s1 = "Это строка."
```

```julia
s2 = """Это тоже строка. """
```

```julia
"Здесь будет "error" потому что не понятно, где начинается какая строка "
```

```julia
"""Гляньте-ка, нет больше "error"!!! """
```

```julia
println("""Еще одна приятная вещь в строковых литералах
    с тройными кавычками заключается в том, что они игнорируют
    начальные отступы, что хорошо для длинных строк в реальном коде.
    Попробуйте здесь чего-нибудь поменять!""")
```

Строки не пишутся с одиночными `'` кавычками - их используют для одиночного символа:

```julia
first(s1)
```

```julia
'⊂'
```

```julia
'Если вы попытаетесь написать строку в одинарных кавычках, вы получите ошибку'
```

## Строковая интерполяция

Вы можете использовать знак доллара в строке, чтобы отобразить выражение Julia внутри строки - либо одну переменную, либо более сложное выражение:

```julia
name = "Вася"
num_fingers = 10
num_toes = 10
println("Привет, я $name.")
println("У меня на руках $num_fingers пальцев и на ногах $num_toes.")
println("Итого $(num_fingers+num_toes) . Согласен, пример так себе...")
```

## Конкатенация строк

Ниже приведены аж три способа объединения строк! <br><br>
Первый способ заключается в использовании функции `string()`. <br>
`string()` преобразует нестроковые входы в строки.

```julia
s3 = "How many cats ";
s4 = "is too many cats?";
котей = 10
```

```julia
string(s3, s4)
```

```julia
string("I don't know, but ", 😺, " is too few.")
```

Также можно использовать `*` для конкатенации!

```julia
s3*s4
```

### Упражнения

#### 2.1 
Создайте строку в которой "hi" сказано 1000 раз, сначала используя `repeat` а потом с помощью возведения в степень, которое использует оператор `*`. Присвойте эту строку переменной `hi`.

<!-- grade_id: cell-5b38bdb1115161b4, points: 0, grade: false, solution: true -->
```julia

```

<!-- grade_id: cell-20b1027a968fbb48, points: 1, grade: true, solution: false -->
```julia
@assert hi == "hihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihi"
```

#### 2.2 
Объявите пару переменных

```julia
a = 3
b = 4
```
и используя их создайте строки:
```julia
"3 + 4"
"7" 
```
и присвойте результаты переменным `c` и `d` соответственно

<!-- grade_id: cell-36beb524c8ecd33b, points: 0, grade: false, solution: true -->
```julia

```

<!-- grade_id: cell-14c3c5b015a97892, points: 1, grade: true, solution: false -->
```julia
@assert c == "3 + 4"
@assert d == "7"
```

```julia

```