# Структуры данных

Как только мы начнем работать со множеством однотипных данных, нам будет удобней хранить данные в структурах, таких как массивы или словари (а не просто создавать множество переменных). 

Типы охватываемых структур данных : 
1. Кортежи 
2. Словари 
3. Массивы

<br>
Кортежи и массивы - это упорядоченные последовательности элементов (поэтому мы можем индексировать их). Словари и массивы являются изменяемыми. 
Мы объясним это подробнее ниже!

## Кортежи (Tuples)

Мы можем создать кортеж, заключив упорядоченную коллекцию элементов в `( )`.

Синтаксис: <br>
```julia
(item1, item2, ...)```

```julia
myfavoriteanimals = ("penguins", "cats", "sugargliders")
```

Мы можем индексировать в этот кортеж,

```julia
myfavoriteanimals[1]
```

но так как кортежи неизменны, мы не можем обновить их

```julia
myfavoriteanimals[1] = "otters"
```

## А теперь в 1.0: Именованные кортежи

Как вы могли догадаться, `NamedTuple`ы это такие же `Tuple`ы за исключением того, что каждый элемент дополнительно имеет имя! У них есть специальный синтаксис, использующий `=` внутри кортежа:

```julia
(name1 = item1, name2 = item2, ...)
```

```julia
myfavoriteanimals = (bird = "penguins", mammal = "cats", marsupial = "sugargliders")
```

Как и обычные `Tuples`, `NamedTuples` упорядочены, так что мы можем получить их элементы с помощью индексации:

```julia
myfavoriteanimals[1]
```

Они также добавляют особую возможность доступа к значениям по их имени:

```julia
myfavoriteanimals.bird
```

## Словари

Если у нас есть наборы данных, связанных друг с другом, мы можем сохранить эти данные в словаре. Мы можем создать словарь, используя функцию `Dict()`.

Syntax:
```julia
Dict(key1 => value1, key2 => value2, ...)```

Хорошим примером является список контактов, где мы связываем имена с номерами телефонов.

```julia
myphonebook = Dict("Jenny" => "867-5309", "Ghostbusters" => "555-2368")
```

В этом примере каждое имя и номер являются парой «ключ» и «значение». Мы можем получить номер Дженни (значение), используя связанный ключ

```julia
myphonebook["Jenny"]
```

Мы можем добавить еще одну запись в этот словарь следующим образом

```julia
myphonebook["Kramer"] = "555-FILK"
```

Давайте посмотрим, как сейчас выглядит наша телефонная книга...

```julia
myphonebook
```

Мы можем удалить Крамера из нашего списка контактов и одновременно получить его номер, используя `pop!`

```julia
pop!(myphonebook, "Kramer")
```

```julia
myphonebook
```

В отличие от кортежей и массивов, словари не упорядочены. Таким образом, мы не можем индексировать их.

```julia
myphonebook[1]
```

В этом примере `julia` думает, что вы пытаетесь получить доступ к значению, связанному с ключом `1`.

### Массивы

Как и в других программах, под массивами в Julia понимается коллекция упорядоченных элементов, размещённая в многомерной сетке. Векторы и матрицы являются частными случаями массивов.


В Julia предусмотрено достаточно много способов создать массив (с учётом типа, значения, размерности и пр.), но по моему скромному мнению в большинстве случаев удобнее пойти путём Matlab-подобных сред и инициализировать массив через перечисление его элементов в квадратных скобках, разделяя строки знаком " ; ". Чтобы обратиться к элементу, укажите его индекс в квадратных скобках, если массив многомерный, перечислите индексы через запятую. ***Индексация, кстати, начинается с единицы.*** Для облегчения доступа индекс последнего элемента хранится в переменной end. 

Вот некоторые из базовых функций для работы с массивами.
+ **length(A)** - число элементов.
+ **ndims(A)** - число размерностей.
+ **size(A)** - кортеж размерностей.
+ **size(A, n)** - размерность в заданном направлении.
+ **copy(A)** - создание копии массива. 

Синтаксис: <br>
```julia
[item1, item2, ...]
```


Например, мы можем создать массив ингридиентов для кекса "Нежность"

```julia
ingridients = ["Meal", "Eag", "Sugar", "Butter", "Milk"]
```

`1` в `Array{String,1}` означает что мы задали одномерный вектор. `Array{String,2}` будет 2d матрицей и т.д. `String` это тип входящих элементов.

Еще мы можем сохранить последовательность чисел

```julia
fibonacci = [1, 1, 2, 3, 5, 8, 13]
```

```julia
mixture = [1, 1, 2, 3, "Пэрри", "Утконос"]
```

Задав массив, мы можем получить отдельные фрагменты данных из этого массива путем индексации в массиве. Например, если мы хотим, получить третий компонент нашего `ingridients`, мы напишем

```julia
ingridients[3]
```

Мы можем использовать индексирование для редактирования существующего элемента массива

```julia
ingridients[3] = "Fish"
```

Да, Julia использует **индексацию начинающуюся с единицы**, а не с нуля как Python.  Войны ведутся и из-за меньших проблем. У меня есть друг мудрый как Соломон, который предлагает раз и навсегда решить эту проблему начиная счет с ½ `:)`

Мы также можем редактировать массив, используя `push!` и `pop!`. `push!` aдобавляет элемент в конец массива, а `pop!` удаляет последний элемент массива. 

Мы можем добавить еще один номер в нашу последовательность Фибоначчи

```julia
push!(fibonacci, 21)
```

а затем удалить его

```julia
pop!(fibonacci)
```

```julia
fibonacci
```

До сих пор я приводил примеры только одномерных массивов скаляров, но массивы могут иметь произвольное количество измерений, а также могут хранить другие массивы.

Например, следующие массивы:

```julia
favorites = [["koobideh", "chocolate", "eggs"],["penguins", "cats", "sugargliders"]]
```

```julia
numbers = [[1, 2, 3], [4, 5], [6, 7, 8, 9]]
```

Ниже приведены примеры двумерных и трехмерных массивов, заполненных случайными значениями:

```julia
rand(4, 3)
```

```julia
rand(4, 3, 2)
```

Будьте осторожны, когда хотите скопировать массивы!

```julia
fibonacci
```

```julia
somenumbers = fibonacci
```

```julia
somenumbers[1] = 404
```

```julia
fibonacci
```

Редактирование `somenumbers` задело `fibonacci`!

В приведенном выше примере мы фактически не делали копию `fibonacci`. Мы только что создали новый способ доступа к записям в массиве, привязанным к `fibonacci`.

Если мы хотим сделать копию массива, привязанного к `fibonacci`, нам следует использовать функцию `copy`.

```julia
# Во первых, восстановим fibonacci
fibonacci[1] = 1
fibonacci
```

```julia
somemorenumbers = copy(fibonacci)
```

```julia
somemorenumbers[1] = 404
```

```julia
fibonacci
```

В этом последнем примере Фибоначчи не был обновлен. Поэтому мы видим, что массивы связаные с `somemorenumbers` и `fibonacci` разные.

### Упражнения

#### 3.1 
Создайте массив, `a_ray`, Спомощью следующего кода:

```julia
a_ray = [1, 2, 3]
```

Добавьте число `4` в конец масива, а затем удалите его.

```julia

```

<!-- grade_id: cell-a6b20b78fbfeb449, points: 1, grade: true, solution: false -->
```julia
@assert a_ray == [1, 2, 3]
```

#### 3.2 
Попробуйте добавить "Emergency" в `myphonebook` задав значение `string(911)`, а потом попробуйте
```julia
myphonebook["Emergency"] = 911
```

Что тут не так?

```julia

```

#### 3.3 
Создайте новый словарь `flexible_phonebook` где номер Дженни будет целочисленным, а у Охотников за привидениями - строковым используя следующий код

```julia
flexible_phonebook = Dict("Jenny" => 8675309, "Ghostbusters" => "555-2368")
```

```julia

```

<!-- grade_id: cell-66f1d636c676c2f2, points: 1, grade: true, solution: false -->
```julia
@assert flexible_phonebook == Dict("Jenny" => 8675309, "Ghostbusters" => "555-2368")
```

#### 3.4 
Добавьте туда ключ "Emergency" со значением `911` (целочисленное).

```julia

```

<!-- grade_id: cell-ee862ab938680259, points: 1, grade: true, solution: false -->
```julia
@assert haskey(flexible_phonebook, "Emergency")
```

<!-- grade_id: cell-c8c40ababe7cdfce, points: 2, grade: true, solution: false -->
```julia
@assert flexible_phonebook["Emergency"] == 911
```

#### 3.5 
Почему в `flexible_phonebook` целочисленное добавляется, а в `myphonebook` нет? Как нам следует инициализировать `myphonebook`, чтоб он мог принимать целочисленные значения?

```julia

```