# Циклы

темы:
1. Цикл `while` 
2. Цикл `for`
<br>

## while

Синтаксис `while` 

```julia
while *условие*
    *тело цикла*
end
```

Например, мы можем использовать `while` чтоб перебрать массив.

```julia
n = 0
while n < 10
    n += 1
    println(n)
end
# то есть, пока n меньше 10ти, прибавлять к нему по одному и распечатывать значение
n
```

```julia
myfriends = ["Ted", "Robyn", "Barney", "Lily", "Marshall"]

i = 1
while i <= length(myfriends)
    friend = myfriends[i]
    println("Hi $friend, it's great to see you!")
    i += 1
end
```

## Цикл for 

Синтаксис `for`

```julia
for *переменная* in *диапазон*
    *тело цикла*
end
```

Мы могли бы использовать цикл `for` для получения тех же результатов, что и в любом из приведенных выше примеров:

```julia
for n in 1:10
    println(n)
end
```

```julia
myfriends = ["Ted", "Robyn", "Barney", "Lily", "Marshall"]

for friend in myfriends
    println("Hi $friend, it's great to see you!")
end
```

Теперь давайте используем циклы `for` для создания таблиц сложения, где значение каждой записи является суммой индексов ее строки и столбца. <br>

Сначала мы инициализируем массив нулями

```julia
m, n = 5, 5
A = fill(0, (m, n))
```

```julia
for i in 1:m
    for j in 1:n
        A[i, j] = i + j
    end
end
A
```

Вот немного синтаксического сахара для того же вложенного цикла `for`

```julia
B = fill(0, (m, n))
```

```julia
for i in 1:m, j in 1:n
    B[i, j] = i + j
end
B
```

Это можно осуществить более стильно используя *array comprehension* (вложение массива).

```julia
C = [i + j for i in 1:m, j in 1:n]
```

### Упражнения

#### 4.1 
Выведите целые числа от 1 до 100 и напечатайте их квадраты.

```julia

```

#### 4.2 
Используя результаты предыдущего задания создайте словарь `squares` который будет содержать целые числа в качестве ключей и квадраты в качестве их пар-значений

```julia
squares[10] == 100
```

```julia

```

<!-- grade_id: cell-54834a2e4e122b70, points: 1, grade: true, solution: false -->
```julia
@assert squares[10] == 100
@assert squares[11] == 121
```

#### 4.3 
Создайте массив `squares_arr` содержащий квадраты всех чисел от 1 до 100.

```julia

```

<!-- grade_id: cell-17d4631c51b8b4f7, points: 1, grade: true, solution: false -->
```julia
@assert length(squares_arr) == 100
@assert sum(squares_arr) == 338350
```