# Функции

Темы:
1. Как объявить функцию
2. Утинная типизация в Julia
3. Mutating vs. non-mutating functions
4. Функции высокого порядка

## Как объявить функцию
Julia дает нам несколько разных способов написать функцию. Первый требует ключевых слов `function` и` end`

```julia
function sayhi(name)
    println("Hi $name, it's great to see you!")
end
```

```julia
function f(x)
    x^2
end
```

Мы можем вызвать любую из этих функций следующим образом:

```julia
sayhi("C-3PO")
```

```julia
f(42)
```

В качестве альтернативы, мы могли бы объявить любую из этих функций в одной строке

```julia
sayhi2(name) = println("Hi $name, it's great to see you!")
```

```julia
f2(x) = x^2
```

```julia
sayhi2("R2D2")
```

```julia
f2(42)
```

Наконец, мы могли бы объявить их как «анонимные» функции

```julia
sayhi3 = name -> println("Hi $name, it's great to see you!")
```

```julia
f3 = x -> x^2
```

```julia
sayhi3("Chewbacca")
```

```julia
f3(42)
```

### Утинная типизация в Julia 

*"Если она крякает как утка, это утка."* 

Функции Джулии будут работать только с теми входными данными, которые имеют смысл. Например, теперь `sayhi` работает над именем этого второстепенного телевизионного персонажа, записанного как целое число ...

```julia
sayhi(55595472)
```

А `f` будет работать и на матрице.

```julia
A = rand(3, 3)
A
```

```julia
f(A)
```

`f` также будет работать со строкой типа "hi", потому что` * `определено для строковых входов как конкатенация строк.

```julia
f("hi")
```

С другой стороны, `f` не будет работать с вектором. В отличие от `A ^ 2`, который является четко определенным, значение` v ^ 2` для вектора, `v`, не является четко определенной алгебраической операцией.

```julia
v = rand(3)
```

```julia
f(v)
```

## Mutating vs. non-mutating functions

По соглашению, функции, сопровождаемые `!`, изменяют свое содержимое, а функции без `!`, не делают этого. 

Например, давайте посмотрим на разницу между `sort` и` sort! `.

```julia
v = [3, 5, 2]
```

```julia
sort(v)
```

```julia
v
```

`sort (v)` возвращает отсортированный массив, который содержит те же элементы, что и `v`, но` v` остается без изменений.

С другой стороны, когда мы запускаем `sort! (v)`, содержимое v сортируется в массиве `v`.

```julia
sort!(v)
```

```julia
v
```

## Функции высокого порядка

### map

`map` - это функция высшего порядка в Julia, которая *принимает функцию* в качестве одного из своих входных аргументов. Затем `map` применяет эту функцию к каждому элементу структуры данных, которую вы передаете. Например, выполнение

```julia
map(f, [1, 2, 3])
```
даст вам выходной массив, где функция `f` была применена ко всем элементам `[1, 2, 3]`
```julia
[f(1), f(2), f(3)]
```

```julia
map(f, [1, 2, 3])
```

Здесь мы возвели в квадрат все элементы вектора `[1, 2, 3]`, а не сам вектор. 

Для этого мы могли бы передать `map` анонимную функцию, а не именованную функцию, такую как

```julia
x -> x^3
```

через

```julia
map(x -> x^3, [1, 2, 3])
```

так возводятся в куб все элементы вектора `[1, 2, 3]`

### broadcast

`broadcast` - это еще одна функция более высокого порядка. `broadcast` - это обобщение` map`, поэтому он может делать все, что может делать `map`, и даже больше. Синтаксис для вызова `broadcast` такой же, как и для вызова` map`

```julia
broadcast(f, [1, 2, 3])
```

и снова мы применили `f` (возведение в кварат) ко всем элементам `[1, 2, 3]`- на этот раз путем "трансляции" `f`. 

Некоторым синтаксическим сахаром для вызова `broadcast` является размещение `.` между именем функции, которую вы хотите от`broadcast`ить, и ее входными аргументами. 

Например, 
```julia 
broadcast (f, [1, 2, 3]) 
``` 
совпадает с 
```julia 
f. ([1, 2, 3]) 
```

```julia
f.([1, 2, 3])
```

Еще раз обратите внимание, насколько это отличается от вызова 
```julia 
f ([1, 2, 3]) 
``` 
Мы можем возвести в квадрат каждый элемент вектора, но мы не можем возвести в квадрат вектор!

Чтобы понять суть, давайте посмотрим на разницу между

```julia
f(A)
```
и
```julia
f.(A)
```
для матрицы `A`:

```julia
A = [i + 3*j for j in 0:2, i in 1:3]
```

```julia
f(A)
```

Как и прежде, мы видим, что для матрицы, `A`,
```
f(A) = A^2 = A * A
``` 

С другой стороны,

```julia
B = f.(A)
```

содержит квадраты всех элементов `A`. 
Этот точечный синтаксис для вещания позволяет нам писать относительно сложные составные поэлементные выражения таким образом, чтобы это выглядело естественным / ближе к математической записи. Например, мы можем написать

```julia
A .+ 2 .* f.(A) ./ A
```

вместо

```julia
broadcast(x -> x + 2 * f(x) / x, A)
```

хотя оба способа отработают одинаково.

### Упражнения

#### 6.1 
Напишите функцию `add_one`, которая добавляет 1 к своему входу.

```julia

```

<!-- grade_id: cell-5119b9e9623c1cb7, points: 1, grade: true, solution: false -->
```julia
@assert add_one(1) == 2
```

<!-- grade_id: cell-50f83d27187a2064, points: 1, grade: true, solution: false -->
```julia
@assert add_one(11) == 12
```

#### 6.2 
Используйте `map` или` broadcast`, чтобы увеличить каждый элемент матрицы `A` на` 1` и присвоить его переменной `A1`.

```julia

```

#### 6.3 
Используйте синтаксис броадкаста с точкой для увеличения каждого элемента матрицы `A1` на `1` и сохранения его в переменной `A2`

```julia

```

<!-- grade_id: cell-f3bd5479679a8fe1, points: 0, grade: true, solution: false -->
```julia
@assert A2 == [3 4 5; 6 7 8;9 10 11]
```