# Множественная диспетчеризация

В этом блокноте мы рассмотрим **множественную диспетчеризацию (мультиметод)**, которая является ключевой особенностью Юлии. Мультиметод делает программное обеспечение *универсальным* и *быстрым*!

#### Начиная со знакомогоr

Чтобы понять множественную диспетчеризацию в Julia, давайте начнем с того, что мы уже видели. 

Мы можем объявить функции в Julia, не предоставляя Julia никакой информации о типах входных аргументов, которые получит функция:

```julia
f(x) = x^2
```

и тогда Юлия сама определит, какие типы входных аргументов имеют смысл, а какие нет:

```julia
f(10)
```

```julia
f([1, 2, 3])
```

#### Настройка типов входных аргументов 

Тем не менее, у нас также есть опции, чтобы явно сообщить Джулии, какие типы разрешены для наших входных аргументов. Например, давайте напишем функцию `foo`, которая принимает только строки в качестве входных данных.

```julia
foo(x::String, y::String) = println("Принял две строки. Попробуй заставить меня их распечатать!")
```

Здесь мы видим, что для ограничения типа `x` и` y` на `String`и (~~не смешно~~), мы просто пишем после имени входного аргумента двойное двоеточие и ключевым словом` String`. Теперь мы увидим, что `foo` работает со `String` и не работает с другими типами входных аргументов.

```julia
foo("hello", "hi!")
```

```julia
foo(3, "поросенка")
```

Чтобы заставить `foo` работать с целочисленными (` Int`) входами, давайте прикрепим `:: Int` к нашим входным аргументам, когда мы объявим` foo`.

```julia
foo(x::Int, y::Int) = println("Принято два целочисленных")
```

```julia
foo(3, 4)
```

Теперь `foo` работает на целых числах! Но посмотрите, `foo` также все еще работает, когда` x` и `y` являются строками!

```julia
foo("hello", "hi!")
```

Вот мы и подступили к сути множественной диспетчеризации. Когда мы объявили

```julia
foo(x::Int, y::Int) = println("Принято два целочисленных")
```
мы не перезаписывали и не заменяли
```julia
foo(y::String, y::String)```

Вместо этого мы просто добавили дополнительный ***метод*** в ***общую функцию*** называющуюся `foo`.

***Общая функция*** это абстрактное понятие, связанное с определенной операцией. 

Например, обобщенная функция `+` представляет концепцию сложения.

***Метод*** конкретная реализация обобщенной функции для *определенных типов аргументов*. 

Например, `+` имеет методы, которые принимают числа с плавающей запятой, целые числа, матрицы и т. д. 

Мы можем использовать `methods()`, чтобы увидеть, сколько методов существует для` foo`.

```julia
methods(foo)
```

Кроме того: как вы думаете, сколько методов существует для сложения?

```julia
methods(+)
```

Итак, теперь мы можем вызывать `foo` для целых чисел или строк. Когда вы вызываете `foo` для определенного набора аргументов, Джулия выведет типы входных данных и отправит соответствующий метод. Это и есть множественная диспетчеризация.

Она делает наш код универсальным и быстрым. Наш код может быть универсальным и гибким, потому что мы можем писать код с точки зрения абстрактных операций, таких как сложение и умножение, а не с точки зрения конкретных реализаций.

В то же время наш код выполняется быстро, потому что Джулия может вызывать эффективные методы для соответствующих типов. Чтобы увидеть, какой метод отправляется при вызове универсальной функции, мы можем использовать макрос which:

```julia
@which foo(3, 4)
```

Давайте посмотрим, что происходит, когда мы используем `@which` с оператором сложения!

```julia
@which 3.0 + 3.0
```

И мы можем продолжать добавлять другие методы в нашу обобщенную функцию `foo`. Давайте добавим тот, который принимает ***абстрактный тип*** `Number`, который включает в себя подтипы, такие как `Int`, `Float64`, и другие объекты, которые вы можете рассматривать как числа:

```julia
foo(x::Number, y::Number) = println("Входные данные числа. Точнее сказать не могу.")
```

Этот метод для `foo` будет работать, например, с числами с плавающей запятой:

```julia
foo(3.0, 4.0)
```

На всякий случай можно добавить, метод с уточной-типизацией `foo`, который будет принимать любые типы:

```julia
foo(x, y) = println("Сам определяй свои типы!")
```

Учитывая методы, которые мы уже написали для `foo`, этот метод будет вызываться всякий раз, когда мы передаем не числа в `foo`:

```julia
v = rand(3)
foo(v, v)
```

### Упражнения

#### 9.1

Расширьте функцию `foo`, добавив метод, который принимает только один входной аргумент, который имеет тип` Bool`, и печатает "foo с одним логическим значением!"

```julia

```

#### 9.2

Убедитесь, что метод, отправляемый при выполнении 
```julia
foo(true) 
```
является тем, который вы написали.

```julia

```

<!-- grade_id: cell-14072e60ae07c1a2, points: 1, grade: true, solution: false -->
```julia
@assert foo(true) == "foo with one boolean!"
```