# Базовая линейная алгебра в Julia
Author: Andreas Noack Jensen (MIT) (http://www.econ.ku.dk/phdstudent/noack/)
(with edits from Jane Herriman)

Сначала давайте определим случайную матрицу

```julia
A = rand(1:4,3,3)
```

и вектор из единиц

```julia
x = fill(1.0, (3,)) # = fill(1.0, 3)
```

Обратите внимание, что $A$ имеет тип `Array {Int64,2}`, но $x$ имеет тип `Array {Float64,1}`. Джулия определяет псевдонимы `Vector {Type} = Array {Type, 1}` и `Matrix {Type} = Array {Type, 2}`. Многие основные операции такие же, как и в других языках

#### Умножение

```julia
b = A*x
```

#### Транспонирование
Как и в других языках, `A'` является сопряженной транспонированной или присоединенной матрицей

```julia
A'
```

что эквивалентно 

```julia
transpose(A)
```

#### Транспонированное умножение

Юлия позволяет нам писать это без *

```julia
A'A
```

#### Решение линейных систем

Задача $Ax = b$ для квадратной матрицы $ A $ решается с помощью функции `\`.

```julia
A\b
```

`A\b` дает нам *минимизацию квадратов*, если мы имеем переопределенную линейную систему («высокая» матрица)

```julia
Atall = rand(3, 2)
```

```julia
Atall\b
```

и *решение минимальных норм наименьших квадратов*, если у нас есть проблема наименьших квадратов с недостатком ранга

```julia
v = rand(3)
rankdef = hcat(v, v)
```

```julia
rankdef\b
```

и * решение минимальных норм наименьших квадратов *, если у нас есть проблема наименьших квадратов с недостатком ранга

```julia
bshort = rand(2)
Ashort = rand(2, 3)
```

```julia
Ashort\bshort
```

# Библиотека LinearAlgebra

Хотя большая часть линейной алгебры доступна в Julia по умолчанию (как показано выше), есть стандартная библиотека с именем `LinearAlgebra`, которая включает в себя гораздо больше релевантных имен и функций. В частности, он предоставляет факторизации и некоторые типы структурированных матриц. Как и во всех пакетах, вы можете использовать эти дополнительные функции в своем сеансе с помощью `using LinearAlgebra`.

### Упражнения

#### 10.1 
Умножьте вектор `v` скалярно сам на себя и сохраните результат в `dot_v`.

```julia
v = [1,2,3]
```

```julia

```

<!-- grade_id: cell-913fef9b0d19cd52, points: 1, grade: true, solution: false -->
```julia
@assert dot_v == 14
```

#### 10.2 
А теперь умножьте `v` матрично на себя (внешнее произведение) присвоив результат переменной `outer_v`

```julia

```

```julia
@assert outer_v == [1 2 3
                    2 4 6
                    3 6 9]
```

<!-- grade_id: cell-e6b6970ffe104df5, points: 1, grade: true, solution: false -->
```julia
@assert cross_v == [0, 0, 0]
```